# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In the /data directory on the Jupyter server, there is a file called `allergies.json` that contains a list of patient allergies.  It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
from pathlib import Path
HOME = str(Path.home())
ALLERGIES_FILE="/data/allergies.json"
```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
### END SOLUTION
```

```python
### BEGIN SOLUTION
def allergy_count(json_file):
    with open(json_file) as f:
        allergies = json.load(f)
        
    return(len(allergies.get('entry')))
### END SOLUTION
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
### BEGIN SOLUTION
### END SOLUTION
```

```python
### BEGIN SOLUTION
def patient_count(json_file):
    patients = set()
    with open(json_file) as f:
        allergies = json.load(f)
        
    for entry in allergies.get('entry'):
        resource = entry.get('resource')
        patient = resource.get('patient')
        name = patient.get('display')
        patients.add(name)
        
    return patients
### END SOLUTION
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut', 'Paul Boal'}
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
### BEGIN SOLUTION
### END SOLUTION
```

```python
### BEGIN SOLUTION
def allergy_per_patient(json_file):
    patients = {}
    with open(json_file) as f:
        allergies = json.load(f)
        
    for entry in allergies.get('entry'):
        resource = entry.get('resource')
        patient = resource.get('patient')
        name = patient.get('display')
        patients[name] = patients.setdefault(name,0) + 1
        
    return patients
### END SOLUTION
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION

def allergy_list(json_file):
    """(json) -> list (nested list)
    Return a list that contain the patient's name, substance they are allergic to, and 
    the allergic reaction (i.e. manifestation due to substance). A list should be created
    for every patient in the given file. 
    """
    
    reactions = []
    with open(json_file) as f:
        allergies = json.load(f)
        
    for entry in allergies.get('entry'):
        patient = entry.get('resource').get('patient').get('display')
        substance = entry.get('resource').get('substance').get('text')
        manifestation = entry.get('resource').get('reaction')[0].get('manifestation')[0].get('text')
        reactions.append([patient,substance,manifestation])
    
    return reactions

### END SOLUTION
```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
output=[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]

assert allergy_list(ALLERGIES_FILE) == output

```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  Solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
import json

### BEGIN SOLUTION

def allergy_reaction(json_file,patient,substance):
    """
    (json,str,str) -> str
    Return a string, which is the allergy manifestation a specific patient will experience based 
    on a specific substance they are allergic to. Check allergy information in the given json file.
    """
    patient_allergies = allergy_list(json_file)
    manifestation = None
    
    for entry in patient_allergies:
        if patient in entry and substance in entry:
            manifestation = entry[2]
            
    return(manifestation)
    
### END SOLUTION
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

---
---

# Stretch (Extra) Problems

Work on either of the stretch problems below can earn you up to 25 free points toward the midterm assignment.  That is, if you complete one of these extra problems successfully, you can skip 1 of the problems that will appear on the midterm exam coming up next week.

The midterm will be distribute this Saturday 3/13.

This assignment is due on Sunday 3/14.  If you are trying for one of these extra problems Slack me, and I'll provide you feedback on how you did on these before end of day Monday 3/15.  That way you can choose what to complete on the midterm.


---
---

### STRETCH for March 2021 - For those looking for an additional challenge

As I've mentioned in class, CMS is now enforcing a rule around price transparency.  Every facility that take Medicare payments is required to publish a "machine readable" file with it's pricing infomration for a number of common procedures across all of the payers they work with.  There are two examples of such files in the `/data/` directory: `whiteriver.json` and `saline.xml`.

If you want to compare contracted prices across these two hospitals, you'll need to read in the information from both of those files into some kind of data structure, then merge the data together from those two files.  See what you can do.

See if you can create an output file that has the following fields:
* HOSPITAL
* PROCEDURE_CODE
* PAYER
* AMOUNT

If you choose to work on this, you may get stuck at some point and you won't know if you're _doing it right_. Make some assumptions. Document your questions in this notebook.



#### 1. Import & Organize White River Data: 

Assumption: If a payer pays $0 for a specific procedure, then they do not have a contract w/ White River. Ignore these payers.


```python
#data = [Hospital,Code,Description,Contracts/Charge]
#data = [str,str,str,[dict]]
#WhiteRiver is list of 100 lists/things (i.e. possible procedures)
#Reorganize Contracts/Charge to match Saline dataset

data = []

import json

with open("/data/whiteriver.json") as WR:
    WRdata = json.load(WR)
    
    for item in WRdata.get('root')['StandardCharges']:
        WRprocedure = [] 
        WRprocedure.append(WRdata.get('root')['HospitalorFacilityName']) #add hospital name in pos. 0
        WRprocedure.append(item['ProcedureCode']) #add code in pos. 1
        WRprocedure.append(item['Description'].strip()) #add name in pos. 2
        
        WRpayers = []
        for entry in item:
            if "_" in entry and item[entry] is not "0":
                WRpayers.append({"Payer":entry,"Charge":item[entry]})  #add payers as a list of dictionaries in pos. 3

        WRprocedure.append(WRpayers)
        data.append(WRprocedure)
            
```

#### 2. Import, Organize, & add Saline Data to WhiteRiver: 

Comment for Paul: Things got difficult/inefficent in the end with the data structure I used for the object data. I couldn't figure out how to merge WhiteRiver and Saline as separate things so I just added Saline to WhiteRiver using the same data structure as WhiteRiver. 

It probably would have made more sense to put/keep procedure names & codes in dictionary format. 

```python
#data = [Hospital,Code,Description,Contracts/Charge]
#data = [str,str,str,[dict]]
#Saline has 1616 lists/things (i.e. possible procedures)
# Add Saline to WhiteRiver; length of data should be 1717

import xml.etree.ElementTree as ET #use Xpath

with open("/data/saline.xml") as S:
    tree=ET.parse(S)
    root=tree.getroot()
    
    hospital_name = root[0].attrib['Name']
    
    for element in tree.iter():
        if element.tag == 'Item':
            procedure_code = element.attrib['Code']

        if element.tag == 'Description':
            procedure_desc = element.text.strip()

        if element.tag == 'Contracts':
            Spayers = []
            
            for child in element.getchildren():
                Spayers.append(child.attrib)
            
            Sprocedure = [hospital_name,procedure_code,procedure_desc,Spayers]
            data.append(Sprocedure)
            
```

```python
len(data)
```

Output:
```
1717
```

```python
data[0]
```

Output:
```
['White River Medical Center',
 '00170',
 'ANESTH PROCEDURE ON MOUTH',
 [{'Payer': 'AETNA_Outpatient', 'Charge': '8400.9904'},
  {'Payer': 'AETNA_EmergencyRoom', 'Charge': '8400.9904'},
  {'Payer': 'AETNAOTHER_Outpatient', 'Charge': '8400.9904'},
  {'Payer': 'AETNAOTHER_EmergencyRoom', 'Charge': '8400.9904'},
  {'Payer': 'QUALCHOICE_Outpatient', 'Charge': '5346.0848'},
  {'Payer': 'QUALCHOICE_EmergencyRoom', 'Charge': '5346.0848'},
  {'Payer': 'WEBTPAAetnaNonemployee_Outpatient', 'Charge': '8400.9904'},
  {'Payer': 'WEBTPAAetnaNonemployee_EmergencyRoom', 'Charge': '8400.9904'}]]
```

#### Notes on json & xml file structures (Ignore)

```python
# How to identify part of a key with different dictionary structures:

# WR_test={"GrossCharge":"1000","Cash":"50","cigna_OPPS":"100","Cigna_Outpatient":"300","Aetna_OPPS":"0","Allwell_EmergencyRoom":"0"}

# payers=[]

# for entry in WR_test:
#     if "_" in entry and WR_test[entry] is not "0":
#         payers.append({"Payer":entry,"Charge":WR_test[entry]})


#test=[{"GrossCharge":"1000"},{"Cash":"50"},{"cigna_OPPS":"100"},{"Cigna_Outpatient":"300"},{"Aetna_OPPS":"0"},{"Allwell_EmergencyRoom":"0"}]

#test_list=[]

#for entry in test:
    #for item in entry.keys():
        #if "_" in item:
            #test_list.append(entry)
            
            
# Working w/ XML file:

#for description in root.iter('Description'):  #loops through all Description 
    #print(description.text) #provides text outside of tag 
    
#tree.findall('.//Contract') #gives all instances of Contract element- use len if needed
#for element in tree.findall('.//Contract') :
    #print(element.attrib) #provides attribute inside tag (i.e. thing associated w/ contract)
    
```

#### 3. Create function w/ procedure code, name, & data as inputs


```python
#Not super efficient/pretty to loop through the object data
#Probably should have used dict to quickly search for code/name function input

def payer_cost(code,data,name = ''):
    """
    (str,list,str) -> list
    Return procedure information (code, name) and negotiated payer prices for 
    procedures done at White River Medical Center and Saline Memorial Hospital. 
    """
    
    index = []

    for cnt,item in enumerate(data):
        if code == item[1] or name == item[2]:
            index.append(cnt)

    if len(index) == 1:
        msg = data[index[0]]
        return(msg)
        
    else:
        msg1 = data[index[0]]
        msg2 = data[index[1]]
        return(msg1,msg2)
         
```

```python
payer_cost('00170',data) #only one hospital has code
#payer_cost('11010',data)  #both hospitals have this code
```

Output:
```
['White River Medical Center',
 '00170',
 'ANESTH PROCEDURE ON MOUTH',
 [{'Payer': 'AETNA_Outpatient', 'Charge': '8400.9904'},
  {'Payer': 'AETNA_EmergencyRoom', 'Charge': '8400.9904'},
  {'Payer': 'AETNAOTHER_Outpatient', 'Charge': '8400.9904'},
  {'Payer': 'AETNAOTHER_EmergencyRoom', 'Charge': '8400.9904'},
  {'Payer': 'QUALCHOICE_Outpatient', 'Charge': '5346.0848'},
  {'Payer': 'QUALCHOICE_EmergencyRoom', 'Charge': '5346.0848'},
  {'Payer': 'WEBTPAAetnaNonemployee_Outpatient', 'Charge': '8400.9904'},
  {'Payer': 'WEBTPAAetnaNonemployee_EmergencyRoom', 'Charge': '8400.9904'}]]
```

---
---

### STRETCH from March 2020 - For those looking for an additional challenge

The Coronavirus is creating quite the stir right now.  There are some sources suggesting that trends show it is going to be significantly more serious than SARS was back in the 2002 timeframe.  Here's one visualization trying to demonstrate that: https://www.reddit.com/r/China_Flu/comments/ev2b4v/i_updated_some_charts_comparing_this_outbreak/

Someone on Kaggle has generously already compiled a dataset based on information from Johns Hopkins about the Coronavirus outbreak.  https://www.kaggle.com/brendaso/2019-coronavirus-dataset-01212020-01262020  Create a Kaggle account, if you don't already have one.  Download this data set and then upload it to your Jupyter Home folder.  (The "up arrow" button is for uploading a file.)

Use Python's built-in `csv` module to read the data from this file and generate the following information: **what are the total confirmed cases in all of Mainland China as of the latest information in the data set?**  Some important things to note:
* Each entry for a given city has the **cumulative** number of cases.  So that column is not additive (it cannot be summed).  You'll have to find a way to filter your data for the last day for each city, then total those up.
* If you choose to parse the date column, you will want to lookup how to do that using Python's `datetime` module.  Especially the `strptime` function.  https://docs.python.org/3/library/datetime.html#strftime-strptime-behavior  Hint: you can parse a date string in the format 2/17/2020 using the code below.  This link will tell you what things like `%m` and `%Y` mean.

```
from datetime import datetime
d = datetime.strptime('2/17/2020', '%m/%d/%Y')
```

If you want to take this another step, **create a list of tuples that contain (observate date, total confirmed) totalled over all locations represented in the data**

```python

```

---

## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Save this note with Ctrl-S (or Cmd-S)
2. Skip down to the last command cell (the one starting with `%%bash`) and run that cell.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

```python
assert False, "DO NOT REMOVE THIS LINE"
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-25-cef8011cb395> in <module>
----> 1 assert False, "DO NOT REMOVE THIS LINE"

AssertionError: DO NOT REMOVE THIS LINE
```

---

```python
%%bash
git pull
git add week06_assignment_2.ipynb
git commit -a -m "Submitting the week 6 programming assignment"
git push
```

Output:
```
Already up to date.
[main 20281f5] Submitting the week 6 programming assignment
 2 files changed, 821 insertions(+), 3 deletions(-)
 create mode 100644 week06/week06_assignment_2.ipynb
To github.com:mmunozru/hds5210-2021.git
   21266e2..20281f5  main -> main
```


---

If the message above says something like _Submitting the week 6 programming assignment_ or _Everything is up to date_, then your work was submitted correctly.

```python

```